Patch-walking: Coordinated multi-pipette patch clamp for efficiently finding synaptic connections

Author response:

The following is the authors’ response to the original reviews.

We thank the reviewers and editors for insightful feedback on how we could improve the manuscript. We have revised the manuscript and addressed the points raised.

Regarding the technical issues raised about the quality of patch clamp recordings (Reviewer 2), we acknowledge that the upper limit of the access resistance cutoff should be lower and that the accepted change should be 10-20%. To this end, we have revised the manuscript to more accurately detail the quality metrics used. The access resistance for the neurons in paired recordings were below 40 MΩ (similar to the metric used by Kolb et al. 2019), and if the access changed above 50 MΩ, we stopped recording from that neuron. Furthermore, the inclusion of neurons in the histogram with access resistance above 50 MΩ was to highlight the total number of neurons patched but not necessarily used in paired recordings. As this was done with an automated robotic system, the neurons would still undergo an initial voltage clamp and current clamp protocol before the pipette would release the neuron and patch another cell. To the point of Reviewer 2, this patch-walk protocol could also be alternatively implemented using manual recording approaches and this point has been included in the revised manuscript.

Regarding the spatial restrictions (Reviewer 3), we agree that the average intersomatic distance is higher than ideal. This was likely due to failed patch attempts; for instance, if one pipette successfully achieved whole cell, and the other pipette had several sequential failed patch attempts, the intersomatic distance (ISD) would increase with each failed attempt due to the user selected index of cells. Ideally, the pipettes would be walking across a slice with low ISD if the whole-cell success rate was closer to 100%. To overcome this challenge in future work, automated cell identification and tracking could enable the path planning to be continuously updated after each patch attempt. Given the whole-cell success rate efficiency for a given electrophysiologist, we believe that the automated robot could be improved in later versions to include routeplanning algorithms to minimize the distance between neurons. Alternatively, this patch-walk system could also be integrated to improve connectivity yields for manual recording approaches as well.

For the point raised about morphological identification, we believe that while important, morphological identification is out of the scope for this project. Future work will include neuronal reconstruction. Regarding the other points, we will amend the manuscript to highlight other key metrics such as maximum time we could hold a neuron under the whole-cell configuration. Additionally, we agree with Reviewer 3 that some of the current language may cause confusion, and we will amend it accordingly.

To all the reviewers, thank you for your time, understanding, and the opportunity to improve our manuscript.

Changes from the earlier version to the revised version

Abstract

Removed in the revised version: Summary

Significant technical challenges exist when measuring synaptic connections between neurons in living brain tissue. The patch clamping technique, when used to probe for synaptic connections, is manually laborious and time-consuming. To improve its efficiency, we pursued another approach: instead of retracting all patch clamping electrodes after each recording attempt, we cleaned just one of them and reused it to obtain another recording while maintaining the others. With one new patch clamp recording attempt, many new connections can be probed. By placing one pipette in front of the others in this way, one can “walk” across the tissue, termed “patch-walking.” We performed 136 patch clamp attempts for two pipettes, achieving 71 successful whole cell recordings (52.2%). Of these, we probed 29 pairs (i.e., 58 bidirectional probed connections) averaging 91 µm intersomatic distance, finding 3 connections. Patch-walking yields 80-92% more probed connections, for experiments with 10-100 cells than the traditional synaptic connection searching method.

Introduction

To elucidate the mechanisms that regulate memory formation, perception, decision-making, and other cognitive functions, scientists seek to measure brain activity at the resolution of individual neurons [1]. However, neurons do not operate in isolation; cognitive function relies on chemical communication between these individual cells, giving rise to neural networks across the entire brain. These connections, or synapses, transmit information and measuring their strength, direction, and other properties is essential to unraveling how the brain works. Yet there are significant challenges that make finding, measuring, and comprehending these dynamic synaptic connections time consuming and low throughput.

Patch clamp recording remains the gold-standard technique for high-fidelity electrophysiological measurements for studying individual neurons and their synaptic connections in the living brain [2, 3, 4]. Patch clamp, with sub-threshold resolution and millisecond precision, has enabled studies ranging from mapping the healthy rodent brain to characterizing the behavior of single cells in neurodegenerative diseases [5, 6, 7]. However, in return for superior signal quality as compared with other methods [8, 9], the traditional patch clamp technique remains low throughput since it is manually laborious and time-consuming [10].

Earlier version: Patch clamp recordings require forming an electrical connection with the intracellular space. This electrical connection requires a high resistance seal between the neuron’s membrane and a clean microelectrode pipette. Detaching the pipette from this seal leaves behind residual membrane material that inhibits the formation a new connection with another neuron. Thus, each pipette must be cleaned or replaced after each recording. When studying cells in intact tissue such as brain slices, skill is needed to avoid neighboring cells and account for tissue deformation. Even skillful efforts only yield around 10 cells recorded per day, and “whole cell” success rates are highly variable, typically ranging from 30%–90% of attempts resulting in successful patch clamp recordings, even for experienced users [11, 12, 10].

Revised version: Whole cell patch clamp recordings require forming an electrical connection between the recording electrode and the membrane of an individual neuron. This electrical connection requires a high resistance seal between the neuron’s membrane and a clean microelectrode pipette. Detaching the pipette from this seal leaves behind residual membrane material that inhibits the formation a new connection with another neuron. Thus, each pipette must be cleaned or replaced after each recording. When studying cells in intact tissue such as brain slices, skill is needed to avoid neighboring cells and account for tissue deformation. Even skillful efforts only yield around 10 cells recorded per day, and “whole cell” success rates are highly variable, typically ranging from 30%–90% of attempts resulting in successful patch clamp recordings, even for experienced users [11, 12, 10].

Removed in the revised version: Scaling traditional, manual patch clamp apparatus to multiple pipettes, in order to obtain synaptic connections between cells, has therefore been extraordinarily technically challenging. The efforts of, for example Perin [13] and Tolias [14] are laudable, but can require years of effort to overcome low throughput and yield. In practice, the approach of Perin et al. and Tolias et al. was to obtain a first patch clamp recording, then shift human attention to a second cell, placing the pipette onto another cell. If a connection was not present, both pipettes were retracted and replaced.

Earlier version: Recently, patch clamp recording efficiency and throughput has increased due to improvements in automated pressure control systems, new algorithms for automated pipette movements guided by visual or electrical signals, and pipette cleaning, rather than reuse [15, 16, 17, 18, 19, 20, 21]. In this way, we have previously developed a robotic system, “the PatcherBot,” capable of performing unattended, multi-hour patch clamp experiments in brain slices, with a whole cell success rate of 51% [22]. These advances rely on the con-current discovery that pipettes can be reused, rather than replaced after each recording attempt [23, 24, 25]. These improvements have enabled novel drug screening assays [26], deep in-vivo recordings [27], voltage indicator screening, and fluorescent cell targeted patch clamp [15].

Revised version: Scaling traditional, manual patch clamp apparatus to multiple pipettes, in order to obtain synaptic connections between cells, has therefore been extraordinarily technically challenging. The efforts of, for example Perin [13] and Tolias [14] are laudable, but can require years of effort to overcome low throughput and yield. Recently, patch clamp recording efficiency and throughput has increased due to improvements in automated pressure control systems, new algorithms for automated pipette movements guided by visual or electrical signals, and pipette cleaning, rather than reuse [15, 16, 17, 18, 19, 20, 21]. In this way, we have previously developed a robotic system, “the PatcherBot,” capable of performing unattended, multi-hour patch clamp experiments in brain slices, with a whole cell success rate of 51% [22]. These advances rely on the concurrent discovery that pipettes can be reused, rather than replaced after each recording attempt [23, 24, 25]. These improvements have enabled novel drug screening assays [26], deep in-vivo recordings [27], voltage indicator screening, and fluorescent cell targeted patch clamp [15].

Earlier version: In the field of connectomics and synaptic physiology, several groups have developed methods for obtaining semi-automated patch clamp recordings of synaptically connected neurons [7, 28, 29, 30, 11]. In the most impressive examples, large-scale connectomics studies have recently emerged from the Allen Institute for Brain Science and Geiger lab. At the Allen Institute, 20,949 connections were probed in the mouse brain [11]. The efficiency of this effort over 1,700 experiments, on average, yielded around 12 potential connections probed per experiment. The Allen Institute leveraged an eight-pipette setup that, on average, successfully connected to four neurons per recording, resulting in an average of 12 possible connections (n2–n = 12). Geiger used the pipette reuse method [23] to increase the number of potential connections probed from 12 to 41 (approximately n=7 cells patched simultaneously) on a comparable eight pipette apparatus [28]. Notably, these papers from the Allen Institute and the Geiger lab used eight manipulators, currently obtainable by only a handful of labs due to complexity and cost.

Revised version: In the field of connectomics and synaptic physiology, several groups have developed methods for obtaining semi-automated patch clamp recordings of synaptically connected neurons [7, 28, 29, 30, 11]. In the most impressive examples, large-scale connectomics studies have recently emerged from the Allen Institute for Brain Science and Geiger lab. At the Allen Institute, 20,949 connections were probed in the mouse brain [11]. The efficiency of this effort over 1,700 experiments, on average, yielded around 12 potential connections probed per experiment. The Allen Institute leveraged an eight-pipette setup that successfully connected to an average of four neurons per recording, resulting in an average of 12 possible connections (n2–n = 12) per experiment. Peng et al. used the pipette reuse method [23] to increase the number of potential connections probed from 12 to 41 (approximately n=7 cells patched simultaneously) on a comparable eight pipette apparatus [28]. Notably, these papers from the Allen Institute and the Geiger lab used eight manipulators, currently obtainable by only a handful of labs due to complexity and cost.

In all previous efforts, both manual and automated, in which multiple pipettes (referred to as multi-patching) are used to probe for synaptic connections, the experimental approach has involved (1) obtaining as many simultaneous recordings as possible, (2) probing their connections, and then (3) retracting all pipettes.

Recognizing how much effort and skill is necessary to obtain many simultaneous recordings, coupled with the advantages of pipette reuse and automation, we hypothesized a novel approach. If instead of retracting all pipettes, perhaps just one of them could be cleaned and reused to obtain a new whole cell recording while maintaining the others. Thus, with one new patch clamp recording attempt, many new connections can be probed. By placing one pipette in front of the others in this way, one can “walk” across the tissue, which we term “patch-walking.” Thus, in this work, we introduce the theory, methods, and experimental results for a fully automated in vitro approach with a coordinated pipette route-planning to “patch-walk” across a brain slice. We demonstrate efficiently recording dozens of neurons using a two-pipette apparatus for finding synaptic connections. Here, we show that this approach, as compared with the traditional approach, increases the rate of potential neurons probed, decreases experimental time, and enables sequential patching of additional neurons.

Results

Mathematical Modeling

Earlier version: The total possible number of connections probed using the traditional method of synaptic patch clamp recording can be expressed as a function of number of cells (n), and number of pipettes (p), as Similarly, the total possible number of connections probed using the patch-walking method can be expressed as To visualize the advantage of patch-walking over the traditional method, these two equations can be represented as a matrix of potential probed connections. For example, the total number of possible connections using the traditional method and patch-walking for a two pipette apparatus is depicted in Fig. 1A and B, respectively. Using these equations, patch-walking is always preferable in practice for n > p. Furthermore, one can expect the improvement in number of connections probed to approach double as n approaches infinity. For practical cases (apparatus with 2-8 pipettes), patch-walking yields 80-92% more probed connections, or efficiency, for experiments with 10-100 cells than the traditional synaptic connection searching method.

Revised version: The total possible number of connections probed using the traditional method of synaptic patch clamp recording can be expressed as a function of number of recorded cells (n), and number of pipettes in the multi-patch apparatus(p), as Similarly, the total possible number of connections probed using the patch-walking method can be expressed as To visualize the advantage of patch-walking over the traditional method, these two equations can be represented as a matrix of potential probed connections. For example, the total number of possible connections using the traditional method and patch-walking for a two pipette apparatus is depicted in Fig. 1A and B, respectively. Using these equations, patch-walking is always preferable in practice for n > p. Furthermore, one can expect the improvement in number of connections probed to approach double as n approaches infinity. For practical cases (apparatus with 2-8 pipettes), patch-walking yields 80-92% more probed connections, or efficiency, for experiments with 10-100 cells than the traditional synaptic connection searching method.

A-B. Schematically, for patch clamp apparatus with two pipettes in search of synaptically connected neurons to record; colored squares represent connections that can be probed using the traditional approach (A) as compared to patch-walking (B). In this schematic, n=8 cells were patched by p=2 pipettes, either in groups of two (A), which yields two possible connections, or by walking across the tissue (B), which yields almost double the number of possible connections. C. A multi-patching apparatus with two pipettes was built with automated pressure control and manipulator movement. D. The software interface used for patch-walking. On the left is the view of the brain slice under the microscope, with the two pipettes highlighted by triangles and user selected cell locations indicated by red circles. On the right are plots used to monitor each step of the patch clamp process: neuron hunting, gigasealing, and membrane test waveform (to monitor break-in state).

Dual-patching experiment

We built the apparatus (Fig. 1C) and developed the software (Fig. 1D) to perform patch-walking with two manipulators. We first conducted a dual-patch throughput experiment for two pipettes patching in a brain slice without testing for connectivity. In 33 patching attempts, (18 attempts for pipette 1 and 15 attempts for pipette 2,) we achieved whole cell success rates for pipette 1 of 44.4% (n=8/18 successful whole cells) and pipette 2 of 46.7% (n=7/15 successful whole cells). This is similar to success rates for manual patching as well as previously reported automated patch clamp robots (43%-51% for Kolb et al. [22]). This result demonstrates the expected throughput and yield of these independent, uncoordinated pipettes.

Earlier version: Next we implemented the coordinated, dual-patching robot. In this series of experiments, we performed patch clamp attempts on 136 cells from 7 animals over a corresponding 7 days, with 2-3 slices per animal. Out of 136 patch clamp attempts for both pipettes, we achieved 130 successful neuron-hunts (95.6%), 71 successful gigaseals (52.2%), and 71 successful whole cell recordings (52.2%). This is again comparable to previously reported auto-mated patch clamp work such as Kolb et al. (51%), Wu et al. (43.2%), and Koos et al. (63.6% for rat, 37% for human) [18, 15]. In addition, our success rates fall within the success range of manual users (30-80%) on differential interference contrast-based patch clamp systems [31]. Thus, coordinating the motion of the pipette via the patch-walking algorithm does not deteriorate the success rate.

Revised version: Next we implemented the coordinated, dual-patching robot. In this series of experiments, we performed patch clamp attempts on 136 cells from 7 animals over a corresponding 7 days, with 2-3 slices per animal. Out of 136 patch clamp attempts for both pipettes, we achieved 71 successful whole cell recordings (52.2%). This is again comparable to previously reported automated patch clamp work such as Kolb et al. (51%), Wu et al. (43.2%), and Koos et al. (63.6% for rat, 37% for human) [18, 15]. In addition, our success rates fall within the success range of manual users (30-80%) on differential interference contrast-based patch clamp systems [31]. Thus, coordinating the motion of the pipette via the patch-walking algorithm does not deteriorate the success rate.

Earlier version: A schematic representation of the experimental arrangement featuring dual patch pipettes in the primary visual and somatosensory cortices is shown in Fig. 2A. In Fig. 2B are histograms showing the distribution of the time it took to achieve a simultaneous recording (left, n = 44), the intersomatic distance between neurons that were patch clamped simultaneously (center, n = 44), and the time to achieve gigaseal for all cells (n = 71). In Fig. 2C are the distributions of whole cell properties (capacitance, tau, input resistance, resting membrane potential, and access resistance) of all cells. We were able to achieve paired patch clamp recordings between two pipettes in an average of 12.6 ± 7.5 min as the pipettes walked across the slice. The average distance between two neurons for screened for connections was 91.6 ± 0.2 µm.

Revised version: A schematic representation of the experimental arrangement featuring dual patch pipettes in the primary visual and somatosensory cortices is shown in Fig. 2A. In Fig. 2B are histograms showing the distribution of the time it took to achieve a simultaneous recording (left, n = 44), the intersomatic distance between neurons that were patch clamped simultaneously (center, n = 44), and the time required to achieve gigaseal (time between increased resistance during neuron hunting step and achieving giga-ohm seal) for all cells (n = 71). In Fig. 2C are the distributions of whole cell properties (capacitance, tau, input resistance, resting membrane potential, and access resistance) of all cells. We were able to achieve paired patch clamp recordings between two pipettes in an average of 12.6 ± 7.5 min as the pipettes walked across the slice. The average distance between two neurons for screened for connections was 91.6 ± 0.2 µm. The cells in paired recordings were held in whole cell configuration up to 45 minutes.

Earlier version: A. An image of a brain slice with a box highlighting the brain region used for experiments: the somatosensory and visual cortices. B. Histograms of patch clamp metrics: time to achieve simultaneous recording (n = 44), distance between neurons during paired recordings (n = 44), and time to achieve gigaseal (n = 71), access resistance (n = 71), and C. Whole cell properties of all cells recorded during patch-walking experiments (n = 71).

Revised version: A. An image of a brain slice with a box highlighting the brain region used for experiments: the somatosensory and visual cortices. B. Histograms of patch clamp metrics: time to achieve simultaneous recording (n = 44), distance between neurons during paired recordings (n = 44), and time to achieve gigaseal (n = 71), and C. Membrane capacitance, time constant (tau), input resistance, resting membrane potential, and access resistance of all cells recorded during patch-walking experiments (n = 71).

Earlier version: We demonstrate a connectivity matrix similar to those done by previous labs such as Peng et al. and the Allen Institute for Brain Science [28, 11]. Of the 71 whole cell recordings we recorded from the robot, we report a yield of 44 paired recordings using our patch-walking technique. In comparison, if we had screened the same 71 neurons for connections using the traditional method, we would have screened for 71/2 = 35 paired recordings. Therefore, our patch-walking algorithm screened for 9 more possible connections for the 71 neurons we tested. Of those 44 recordings, 29 paired recordings (i.e. 58 probed connections) passed quality control checks and were used to validate the efficiency of the patch-walking algorithm. From Eqs. 1 and 2, we see the additional connections screened with the effect of patch-walking.

Revised version: We demonstrate a connectivity matrix similar to those done by previous labs such as Peng et al. and the Allen Institute for Brain Science [28, 11]. Of the 71 whole cell recordings we recorded from the robot, we report a yield of 44 paired recordings using our patch-walking technique. In comparison, if we had screened the same 71 neurons for connections using the traditional method, we would have screened for 71/2 = 35 paired recordings. Therefore, our patch-walking algorithm screened for 9 more possible paired recordings for the 71 neurons we tested. Of those 44 recordings, 29 paired recordings (i.e. 58 probed connections) passed quality control checks and were used to validate the efficiency of the patch-walking algorithm, resulting in 3 found synaptic connections. From Eqs. 1 and 2, we see the additional connections screened with the effect of patch-walking.

According to Perin et al [13], at an intersomatic distance of 91.6 ± 0.171 µm, the expected connection probability is 16.9% for each paired recording. Assuming they are independent, we would expect greater than 50% probability of getting at least 1 connection after just 3 paired recordings. According to binomial probability theory, we had a probability of 89% to find 3 connections with 29 paired recordings.

Earlier version: Fig. 3 shows a connectivity matrix (as in Fig. 1B), a spatial representation of the cells patched cells and connection probed, as well as a representative connection found between two cells. The matrix in Fig. 3A shows the whole cell current clamp protocol described previously (black traces in leftmost column). During paired recordings, one cell would be stimulated in current clamp (traces along the diagonal). Recording color corresponds to a pair of cells tested for connectivity as in Fig. 1, where each color has two traces because each pair of cells can be connected bidirectionally. The nomenclature for each row and column is n.p where n represents the cell number, in this case ranging from 1-7, since 7 total cells were patched between both pipettes, and p labeled either a or b represents each of the two pipettes.

Revised version: Fig. 3 shows a connectivity matrix (as in Fig. 1B), a spatial representation of the cells patched cells and connection probed, as well as a representative connection found between two cells. The matrix in Fig. 3A shows the whole cell current clamp protocol described previously (black traces in leftmost column). During paired recordings, one cell would be stimulated in current clamp (traces along the diagonal). Recording color corresponds to a pair of cells tested for connectivity as in Fig. 1, where each color has two traces because each pair of cells can be connected bidirectionally. The nomenclature for each row and column is n.p where n represents the cell number, in this case ranging from 1-7, since 7 total cells were patched between both pipettes, and p labeled either a or b represents each of the two pipettes. Cell 7 exhibited signs of decreased cell health, likely due to the duration of the experiment and increasing physical disruptions to the slice during patch-walking. The representative connection shown in more detail in Fig 3C was found between cells 1 and 2, with pre-synaptic cell 1 (black) stimulated and cell 2 (red) recording in voltage clamp the post synaptic currents elicited.

A. Matrix of voltage and current traces from 7 neurons in one acute brain slice recorded using the patch-walking algorithm for the robot. Left column shows the firing pattern of the recorded neurons. Cells are numbered such that the number represents the cell and the letter represents the manipulator (a or b). Scale bars: Horizontal 200 ms for firing pattern and connection screening. Vertical 40 mV for action potentials, 50 pA for postsynaptic traces. B. Patch-walking scheme of all neurons from the experiment matrix in (a). The curved lines between neurons represent probed connections in the matrix in (A). C. The probed connection from the connectivity matrix in (A). The stimulus was sent to cell 1 (black) and the response from cell 2 (red) was recorded and averaged over three sweeps.

Removed in the revised version: The representative connection shown in more detail in Fig 3C was found between cells 1 and 2, with pre-synaptic cell 1 (black) stimulated and cell 2 (red) recording in voltage clamp the post synaptic currents elicited.

Discussion

Earlier version: We introduce a variation on the multi-patching technique which we termed patch-walking. Patch-walking enables, theoretically, almost twice the number of connections to be probed on a patch clamping apparatus for a given number of cells patched and pipettes on the rig. We recommend this method for those searching for local synaptic connections using a rig with a small number (e.g, two) of pipettes, such as we have shown. Furthermore, patch-walking provides other advantages: it damages the tissue less between probed connections, decreasing the chance of severed connections, it saves time between experiments since only one pipette is moved at a time, and it enables more recordings from a tissue before cell death which is advantageous for studying rare tissues such as human brain samples.

Revised version: We introduce a variation on the multi-patching technique which we termed patch-walking. Patch-walking enables, theoretically, almost twice the number of connections to be probed on a patch clamping apparatus for a given number of cells patched and pipettes on the rig. We recommend this method for those searching for local synaptic connections using an apparatus with a small number (e.g, two) of pipettes, such as we have shown. Additionally, patch-walking causes less tissue damage because it only requires one new pipette to enter the slice for each connection being probed, compared to the traditional method that needs two new pipettes for each connection. Further, this method saves time between probed connections since only one pipette is moved at a time and it enables more recordings from a tissue before cell death which is advantageous for studying rare tissues such as human brain samples.

For scaling patch-walking beyond the apparatus described here, we caution that one must take into account pipette collisions and choose cell-pipette assignments carefully. Future work to improve upon patch-walking could include developing an optimal route-planning path such as the Monte Carlo Tree algorithm as a strategy to optimize the pipette-cell assignments, considering (1) spacing pipettes apart in order to avoid collisions between the fragile glass pipettes while (2) maximizing the probability of a connection. Specifically, we would recommend a threshold of inter-cell distance to be less than 200 µm in order to have at least a 10% probability of connection according to Perin et al. [13].

The patch-walking algorithm can make multiple-pipette patch-clamp electrophysiology more accessible to a wider range of laboratories that usually conduct simultaneous recordings with several manipulators. For example, conducted studies such as Galarreta et al. studying a network of parvalbumin fast-spiking GABAergic interneurons could use this technology to make them more efficient to days worth of experiments rather than months [32]. Further, the robot could also be altered to include users in the loop if they want to have control over certain aspects of the patching process or enable experienced patchers with digital pressure control. Even the best human electrophysiologists can only control one manipulator at a time, but the robot can control multiple pipettes, pressure regulators, and command signals independently. Patch-walking offers throughput improvements over manual patching, especially for those looking to utilize paired recordings in their experiments [33, 34, 35, 36].

Removed in the revised version: Limitations of the study

Out of the 29 paired recordings, we found 3 synaptic connections. While this number of connections is lower than predicted according to Perin et al. [29] based on the intersomatic distances between these cells, we hypothesize that this is most likely due to biological variation.

While the patch-walking method provides an efficient means to probe for synaptic connections using two pipettes, introducing additional pipettes presents notable challenges. Specifically, as pipettes maneuver into and out of the brain slice, tissue deformation often disturbs any pipettes in whole cell configuration.

Earlier version: Future applications and variations in patch-walking could include the use of channelrhodopsin-assisted circuit mapping [37, 38] to enable larger circuit mapping with multiple patch electrodes. Patch-walking could also be used for fluorescent-targeted cells wherein one pipette could target a specific subset of cells while the other pipette would probe off-target cells. A third alternative could be that one pipette patches a deep cell and stays patched onto it while secondary pipettes continue to automatically patch other cells and search for connections. Experiments could be done to probe specific layer to layer connections, and machine learning algorithms to detect specific neurons could be introduced as well [39]. From these potential uses of patch-walking, future work of automated multi-patching could greatly enhance the different methods used to study functional synaptic connectivity and local neuronal networks.

Revised version: Future applications and variations in patch-walking could include the use of channelrhodopsin-assisted circuit mapping [37, 38] to enable larger circuit mapping with multiple patch electrodes. Patch-walking could also be used for fluorescent-targeted cells wherein one pipette could target a specific subset of cells while the other pipette would probe off-target cells. A third alternative could be that one pipette patches a deep cell and stays patched onto it while secondary pipettes continue to automatically patch other cells and search for connections. Additionally, this patch-walk protocol could also be implemented into manual recording approaches, leveraging the idea that only one pipette has to patch onto a new cell to test for connections, as opposed to two pipettes. Future work can include morphological identification or layer-to-layer connectivity studies. Further, machine learning algorithms to detect specific neuronal subtypes could be integrated for improved, real-time route-planning [39]. From the presented methodology of patch-walking and potential future applications, patch-walking can be a useful tool to study synaptic connectivity, especially for researchers new to the field of single cell electrophysiology.

Methods

Automated patch clamp apparatus

We designed and implemented an experimental apparatus to demonstrate the utility of patch-walking. The apparatus features a standard electrophysiology rig with two PatchStar micromanipulators. Samples (mouse brain slices) were imaged using a 40X objective (LUMPLFL40XW/IR, NA 0.8, Olympus) on a motorized focus drive, illuminated under differential interference contrast microscopy (DIC) with an infrared light-emitting diode (780 nm), and captured with a Rolera Bolt camera (QImaging). We used a peristaltic pump (120S/DV, Watson-Marlow) to perfuse the brain slices with buffer solution. We utilized the brain slice sample holder with integrated cleaning and rinse solution chambers as described previously [23]. We followed the cleaning protocol as suggested by Kolb et al [23], however we did not include rinsing in the cleaning protocol because recent literature found that there is no impediment to the whole cell yield or quality of recording [24, 28].

Electrode pressure was controlled using a custom pipette pressure controller enabled up to four-channels, adapted from prior work [22, 16]. Briefly, for each pipette, pressure was controlled by a ±10 psi regulator (Proportion-Air) using an analog (0-10 V) control signal. The control signal for each regulator was generated by a microcontroller (Arduino Due) via a digital-to-analog converter (MAX539, Maxim Integrated). In order to minimize valve switching to efficiently scale up the patcherBot to multiple manipulators and pipette pressure control, a custom printed circuit board was developed to control up to a maximum of four manipulators. Individual pressure regulators for each pipette were necessary to ensure that different pressures could be maintained on each pipette. The custom pressure controller regulates house-air line to deliver −500 to +700 mbar (relative to sea-level) using an inline venturi tube (SMC) and solenoid valve (Parker Hannifin) for rapid pressure switching [40, 22, 17, 26].

Earlier version: For real-time electrophsyiology feedback and collection, we used the Multiclamp 700B amplifier (Molecular Devices), cDAQ-9263/9201 and USB-6221 OEM data acquisition boards (National Instruments), and Axon Digidata 1550B. The two data acquisition boards were each assigned to a manipulator in order to simultaneously acquire different signals from each pipette. This is particularly important for asynchronous and independent pipette control. Machine vision-based pipette calibration and position correction was performed according to Gonzalez et al. [41] to correct for small micromanipulator position errors.

Revised version: For real-time electrophsyiology feedback and collection, we used the Multiclamp 700B amplifier (Molecular Devices), cDAQ-9263/9201 and USB-6221 OEM data acquisition boards (National Instruments), and Axon Digi-data 1550B. The two data acquisition boards were each assigned to a manipulator in order to simultaneously acquire different signals from each pipette. This is particularly important for asynchronous and independent pipette control. Machine vision-based pipette calibration and position correction was performed according to Gonzalez et al. [41] to correct for small micromanipulator position errors.

Following brain slice preparation and pipette fabrication, filling, and installation, pipette location was calibrated according to Kolb et al [22] for all pipettes, resulting in a ‘home’ position as described for each pipette.

Brain slice preparation

Earlier version: All animal procedures were in accordance with the US National Institutes of Health Guide for the Care and Use of Laboratory Animals and were approved by the Institutional Animal Care and Use Committee at the Georgia Institute of Technology. For the brain slice experiments, male mice (C57BL/6, P19–P36, Charles River) were anesthetized with isofluorane, and the brain was quickly removed. Coronal sections (300 µm thick) were then sliced on a vibratome (Leica Biosystems VT1200S) while the brain was submerged in ice-cold sucrose solution containing (in mM) 40 NaCl, 4 KCl, 1.25 NaH2PO4·H2O, 7 MgCl2, 25 NaHCO3, 10 D-Gluocse, 0.5 CaCl2·2H2O, 150 Sucrose (pH 7.3–7.4, 300–310 mOsm). The slices were incubated at 37 °C for 1 h in neuronal artificial cerebro-spinal fluid (aCSF) consisting of (in mM) 124 NaCl, 2.5 KCl, 1.25 NaH2PO4·H2O, 1.3 MgCl2, 26 NaHCO3, 10 D-Gluocse, 2 CaCl2·2H2O, 1 L-Ascorbate·H2O (pH 7.3–7.4, 290–300 mOsm). Prior to recording, the slices were maintained at room temperature for at least 15 min (22 °C–25 °C). The sucrose solution and neuronal ACSF were bubbled with 95% O2/5% CO2. Recordings were performed in mouse primary visual area and somatosensory cortex.

Revised version: All animal procedures were in accordance with the US National Institutes of Health Guide for the Care and Use of Laboratory Animals and were approved by the Institutional Animal Care and Use Committee at the Georgia Institute of Technology. For the brain slice experiments, male mice (C57BL/6, P19–P36, Charles River) were anesthetized with isofluorane, and the brain was quickly removed. Coronal sections (300 µm thick) were then sliced on a vibratome (Leica Biosystems VT1200S) while the brain was sub-merged in ice-cold sucrose solution containing (in mM) 40 NaCl, 4 KCl, 1.25 NaH2PO4·H2O, 7 MgCl2, 25 NaHCO3, 10 D-Gluocse, 0.5 CaCl2·2H2O, 150 Sucrose (pH 7.3–7.4, 300–310 mOsm). The slices were incubated at 37 °C for 1 h in neuronal artificial cerebro-spinal fluid (aCSF) consisting of (in mM) 124 NaCl, 2.5 KCl, 1.25 NaH2PO4·H2O, 1.3 MgCl2, 26 NaHCO3, 10 D-Gluocse, 2 CaCl2·2H2O, 1 L-Ascorbate·H2O (pH 7.3–7.4, 290–300 mOsm). Prior to recording, the slices were maintained at room temperature for at least 15 min (22 °C–25 °C). The sucrose solution and neuronal ACSF were bubbled with 95% O2/5% CO2. Recordings were performed in mouse primary visual area and somatosensory cortex.

Patch-clamp recording

Borosilicate pipettes were pulled on the day of the experiment using a horizontal puller (P-1000, Sutter Instruments) to a resistance of 4–6 MΩ. The intracellular solution was composed of (in mM) 135 K-Gluconate, 10 HEPES, 4 KCl, 1 EGTA, 0.3 Na-GTP, 4 Mg-ATP, 10 Na2-phosphocreatine (pH: 7.2–7.3, 290–300 mOsm). Recordings were performed at room temperature with constant perfusion of oxygenated neuronal aCSF. Pipette pressure during patch clamp steps was digitally controlled and pipettes were cleaned according to Kolb et al [23, 22], as previously described.

Patch-walking experimental method

The patch-walking algorithm is depicted schematically in Fig. 4. At the beginning of each patch-walking experiment, cells are first selected by the user. For each slice experiment, we selected 8-10 healthy cells located 20-100 um below the surface of the tissue. These cells were spread across an area of approximately 200 um x 200 um (note the field of view under 40x magnification is approximately 50 um). These cell locations with three dimensional coordinates are stored in a cell queue for subsequent patch attempts. From these cell coordinates, the robot computes the distance between the pipettes’ respective home positions and each cell.

Schematic of the patch-walking experimental workflow. The patch-walking process begins with selection of cells by the user. The Patcherbot then assigns cells to each pipette based on their distance to the pipettes’ home positions. Each pipette works in parallel, only working independently during steps which require the camera and stage (ie neuron hunting, neuron detection). Once a pipette has achieved these steps successfully, the stage and camera are designated to the other pipette. If the pipette failed the patch attempt, it is cleaned and reused. Once both pipettes achieved whole cell configuration, they are tested for synaptic connectivity. In order to ”patch-walk,” the first pipette to achieve whole cell is released to clean and obtain a new whole cell recording.

The assignment of cells to pipettes proceeds as follows. When a pipette is available for a patch clamp attempt, the cell with the shortest distance to the pipette home position is removed from the cell queue and assigned to the pipette. Thus each pipette initially attempts to patch its closest cell. To probe connections between neurons, both pipettes must form a successful connection. Therefore, in the event that an attempted connection fails, the pipette is retracted and the next closest cell is assigned from the queue. This process is repeated until both pipettes are simultaneously connected to neurons. At this moment, the connections can be tested as described in ”connectivity testing.” A pipette that achieves whole cell configuration is held while patch attempts are made with the other pipette. Unsuccessful patch attempts result in the pipette being cleaned [22], assigned a new cell and then a new attempt. For each patch clamp attempt, control of the microscope objective is assigned to the pipette in the ”pipette finding” or ”neuron hunting” phase, until the pipette has successfully completed “neuron hunting” (when the pipette resistance increases by 0.2 MΩ over 5 descending 0.1 µm steps) [16]. Once both pipettes have established whole cell patch clamp recordings, the connection test is performed. This algorithm then repeated the process until all viable neurons had been patched.

The “patch cell” step in Fig. 1 includes the following: neuron hunting, neuron detection, gigasealing, break in, whole cell protocol. As in Kolb et al [22], and briefly restated here, once the measured resistance reaches 1 GΩ, the algorithm waits (5 s) and proceeds to the break-in state. Break-in is accomplished by short pulses of suction (100-1000 ms, −345 mBar). A break-in is considered successful when the measured resistance drops to under 800 MΩ and the holding current remains low (< −200 pA at −70 mV in slices). The whole cell electrophysiology protocol consists of a voltage clamp protocol where cell parameters (access resistance, membrane resistance, holding current) are measured as well as a current clamp protocol (0 pA for 1s, −300 to +300 pA step for 1s, 0 pA for 1 s). Injected current pulses were 3 second pulses from −20pA to +280 pA in 20 pA steps with a 2 sec, −20 pA hyperpolarizing step 500 ms prior.

We define a paired recording as a pair of cells which are simultaneously patch clamped using the patch-walking experimental method. We define probed connections as twice the number of paired recordings, because the connections can be bi-directional. We define possible connections as the theoretical upper limit of probed connections given a number of pipettes and number of cells recorded. Probed connections is, practically speaking, less than possible connections since patch clamp yield is 30-80% based on our experience with the PatcherBot, depending on sample preparation and tissue and cell type.

Recording quality criteria

Earlier version: All recordings that were retained had access resistance less than 50 MΩ. If the seal quality decreased during recording, the cell is excluded from analysis.

Revised version: The access resistance for the neurons in paired recordings were below 40 MΩ, similar to the metric used by Kolb et al [22], and if the access changed above 50 MΩ, we stopped recording from that neuron. If the seal quality decreased during recording, the cell is excluded from analysis.

Connectivity Testing

We tested for connectivity in a manner similar to that done previously by Peng et al [13] and the Allen Institute for Brain Science [11]. To perform connectivity testing between two simultaneously patched neurons, we performed the following procedure. The BNC cables were manually moved from the NI DAQ to the Digitizer to enable Clampex control of the cells (rather than LabView). Two protocols were run in order to test for the two possible directions of connectivity. For each protocol, one pipette sent a stimulus in current clamp mode to elicit five action potentials at 20 Hz while the other pipette holding in voltage clamp recorded post-synaptic currents, held at −70 mV. Following this bi-directional measurement, the BNC cables were replaced manually to resume patch-walking.

Added in the revised version: Author Contributions

Added in the revised version: C.R.F. and M.C.Y. conceptualized the method. M.C.Y., C.R.L., C.F.L., I.K., W.M.S., and M.M.G. designed experiments; M.C.Y. and M.M.G. performed experiments; M.C.Y and M.M.G performed analyses; M.F. contributed to the interpretation of the results. M.J.M.R., E.S.B., and C.R.F. supervised the research. M.C.Y., M.M.G., and C.R.F. wrote the manuscript. C.R.F. secured funding. All authors provided feedback and helped shape the research, analysis, and manuscript.

Added in the revised version: Declaration of Interests

Added in the revised version: I.K., W.M.S, and C.R.F. are co-inventors on a patent describing pipette cleaning that is licensed by Sensapex.